// tests1.rs
//
// 测试对于确保您的代码按照您的预期工作非常重要。可以使用以下命令在此文件上运行测试：
// rustlings run tests1
//
// 这个测试有一个问题 -- 让测试能够编译！让测试通过！再让测试失败！
//
// Execute `rustlings hint tests1` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

#[cfg(test)]
mod tests {
    #[test]
    fn you_can_assert() {
        assert!(true);
    }
}
